<problem url_name="Activity_1_0">
  <script display_name="Activity 1" type="loncapa/python" url_name="Activity_1_1">

# This Python code defines and evaluates the problem. You usually don't need to touch the
# first portion of the code and can jump to "Problem definition" near the end.
import json
class ResponseCell(object):
    """Abstract base class for response cells."""
    TYPE = 0
    def get_state(self):
        """Return a dictionary describing the instance for sending to the JS code."""
        d = vars(self).copy()
        d.update(type=self.TYPE)
        return d
class Numeric(ResponseCell):
    """A numeric response cell."""
    TYPE = 1
    def __init__(self, answer, tolerance=None):
        """Set the correct answer and the allowed relative tolerance in percent."""
        self.answer = answer
        if tolerance is None:
            tolerance = default_tolerance
        self.abs_tolerance = abs(answer) * tolerance / 100.0
    def check_response(self, student_response):
        """Return a Boolean value indicating whether the student response is correct."""
        try:
            r = float(student_response)
        except ValueError:
            return False
        return abs(r - self.answer) &lt;= self.abs_tolerance
class String(ResponseCell):
    """A string response cell."""
    TYPE = 2
    def __init__(self, answer):
        """Set the correct answer."""
        self.answer = answer
    def check_response(self, student_response):
        """Return a Boolean value indicating whether the student response is correct."""
        return student_response == self.answer
def init(table, help_text, column_widths, row_height):
    """Initialisation code for values exported to the XML code.
    The initial_state is a JSON string exported to the JavaScript code and used to construct
    the table.  The height is the height of the iframe in pixels.
    """
    data = [
        [cell.get_state() if isinstance(cell, ResponseCell) else str(cell) for cell in row]
        for row in table
    ]
    initial_state = dict(
        data=data,
        help_text=help_text,
        column_widths=column_widths,
        row_height=row_height,
    )
    # Curiously, the JSON string must be doubly HTML-escaped.
    initial_state = json.dumps(initial_state).replace('&amp;', '&amp;amp;amp;').replace('"', '&amp;amp;quot;')
    # All heights and widths include a bit of safety margin to avoid scrollbars and rendering
    # artifacts on some browsers.
    if column_widths is not None:
        width = sum(column_widths) + 2
    else:
        width = 800 + 2
    height = (row_height + 1) * len(table)
    if help_text:
        height += (len(help_text) // 100 + 2) * 26
    return initial_state, width, height
def check_consistency(state):
    """Check whether the state of the table is consistent with the initial state.
    If not, the problem was edited in the mean time, and we ask the student to click
    the reset button (as we can't initiate resetting the problem from code).
    """
    if len(table) != len(state['data']):
        return False
    for row, row_state in zip(table, state['data']):
        if len(row) != len(row_state):
            return False
        for cell, cell_state in zip(row, row_state):
            if isinstance(cell, ResponseCell):
                if not isinstance(cell_state, dict):
                    return False
                orig_state = cell.get_state()
                if not all(orig_state[k] == cell_state[k] for k in orig_state):
                    return False
            else:
                if not str(cell) == cell_state:
                    return False
    return True
def check_table(unused_expect, ans):
    """Determine the correctness of all student responses.
    The correctness is independently determined by the JavaScript code on the client side.
    """
    ans = json.loads(ans)
    state = json.loads(ans['state'])
    if not check_consistency(state):
        return {
            'ok': False,
            'msg': 'The problem was edited by the course staff since the last time you '
            'worked on it.  Please click the "Reset" button to get the latest version.'
        }
    correct_answers = total_questions = 0
    for row, row_state in zip(table, state['data']):
        for cell, cell_state in zip(row, row_state):
            if isinstance(cell, ResponseCell):
                total_questions += 1
                correct_answers += cell.check_response(cell_state['value'])
    if correct_answers == total_questions:
        return {'ok': True, 'msg': 'Great job!'}
    else:
        return {
            'ok': False,
            'msg': 'You have {} cells out of {} correct.'.format(correct_answers, total_questions)
        }
## Problem definition ##
## Edit the table and settings below. ##
# The default tolerance in percent, used for all numeric fields that don't specify an ecplicit
# tolerance.
default_tolerance = 0.1
# The table itself.  The first row contains the column headers.  Any cell in the following rows can
# be either a string or number, which will simply be shown to the student, or one of the special
# cell types Numeric() or String(), indicating a numeric or string input field that must be filled
# in by the student.
# Syntax of the special cell types:
#     Numeric(answer=&lt;correct_answer&gt;, tolerance=&lt;tolerance in percent&gt;)
#     String(answer='&lt;correct answer&gt;')
table = [
    ['Year','$ Annual Avg. Price','Price Index (Base Year 2000)','% Inflation Rate (vs prior year)'],
    ['2000', 1.51, Numeric(answer=1.00, tolerance = 3.0),''],
    ['2001', 1.46, Numeric(answer=0.97, tolerance = 3.0),Numeric(answer=-3, tolerance = 10.0)],
    ['2002', 1.36, Numeric(answer=0.90, tolerance = 3.0),Numeric(answer=-7, tolerance = 10.0)],
    ['2003', 1.59, Numeric(answer=1.05, tolerance = 3.0),Numeric(answer=17, tolerance = 10.0)],
    ['2004', 1.88, Numeric(answer=1.24, tolerance = 3.0),Numeric(answer=18, tolerance = 10.0)],
    ['2005', 2.30, Numeric(answer=1.52, tolerance = 3.0),Numeric(answer=22, tolerance = 10.0)],
    ['2006', 2.59, Numeric(answer=1.71, tolerance = 3.0),Numeric(answer=13, tolerance = 10.0)],
    ['2007', 2.80, Numeric(answer=1.85, tolerance = 3.0),Numeric(answer=8, tolerance = 10.0)],
    ['2008', 3.27, Numeric(answer=2.17, tolerance = 3.0),Numeric(answer=17, tolerance = 10.0)],
    ['2009', 2.35, Numeric(answer=1.56, tolerance = 3.0),Numeric(answer=-28, tolerance = 10.0)],
    ['2010', 2.79, Numeric(answer=1.85, tolerance = 3.0),Numeric(answer=19, tolerance = 10.0)],
    ['2011', 3.53, Numeric(answer=2.33, tolerance = 3.0),Numeric(answer=27, tolerance = 10.0)],
    
]

# Set the text that gets displayed when clicking the "+help" button.  Use "None" to disable the
# help feature.
#help_text = None
help_text = "Fill in the yellow fields with the right values."
# Set the width of the columns in pixels.  The total width of the table shouldn't be more than 800.
# You can use "column_width = None" for equal-width columns with a total width of 800 pixels.
column_widths = None
#column_widths = [150, 150, 150, 150]
# The height of a table row in pixels.  Rows may grow higher than the specified value if the text
# in some cells in the row is long enough to get wrapped in more than more lines.  If this happens,
# you might need to replace "$height" close to the end of this file by a big enough value to
# prevent a vertical scrollbar for the iframe from appearing.
row_height = 36
# Initialisation -- don't remove!
initial_state, width, height = init(table, help_text, column_widths, row_height)

</script>
  <p display_name="Activity 1" url_name="Activity_1_2">Use the information in the table below to practice price index and inflation calculations. This example involves calculation of a simple price index based on gas prices. Remember that any price index can be calculated as:</p>
  <p display_name="Activity 1" url_name="Activity_1_3">(Cost of market basket in given year)  /  (Cost of market basket in base year)</p>
  <p display_name="Activity 1" url_name="Activity_1_4">And, the inflation rate is simply the percent change in the price index from one year to the next.  Be sure to round your answer to the nearest whole percent. </p>
  <p display_name="Activity 1" url_name="Activity_1_5">
    <i>Note, all input values should use decimal notation where appropriate. For example, 1/10 should be input as 0.1.</i>
  </p>
  <customresponse cfn="check_table" display_name="Activity 1" url_name="Activity_1_6">
    <jsinput get_statefn="ActiveTable.getState" gradefn="ActiveTable.grade" height="$height" html_file="/static/activetable.html" initial_state="$initial_state" set_statefn="ActiveTable.setState" sop="true" width="$width"/>
  </customresponse>
  
  <p><i>This problem has been updated since launch. You may need to click the reset button before attempting the problem. We recommend checking a single cell before entering all the data.</i></p>
  
  <solution display_name="Activity 1" url_name="Activity_1_7">
    <img src="/static/Screen_Shot_2015-07-14_at_10.50.49_AM.png"/>
  </solution>
</problem>
